# Exercise Sheet 4
## Instructions
The instructions for this exercise sheet are the same as the previous, so this section is just a summary to serve as a reminder. Refer to the material on Engage and at the start of Exercise Sheet 1 for more detail, and if you are unsure ask on the Q&A forum.

This exercise sheet counts towards your overall grade for the unit.

Complete each question in the code cell provided, which will usually contain some skeleton code to help you start the question. Unless specified otherwise you may change the skeleton code as you wish, but your code must pass the formatting tests in the following code cell – if the cell runs without errors then your code is eligible for submission. However these will only test the basics of the question. Your code will be subject to additional hidden tests which will determine your grade. Make sure you thoroughly test your own code to check you think it works correctly before submission.

## Questions
### Question 1 & 2
*(worth 25% each)*

In the cell below, implement 
1. the *bubble sort* and 
2. the *merge sort* algorithms.

Unlike the examples in the unit material, please make your functions *return new sorted versions* of the input lists, *without* modifying the input lists. In other words, do not worry about making your algorithms in-place.

If you find it easier to think about modifying the list in-place (more likely with bubble sort than merge sort), you can always write `new_list = in_list.copy()` on the first line, sort `new_list` in place, then return `new_list`.

#### Bubble Sort
Bubble sort works by iterating over the list, comparing adjacent items, and swapping them if they are out of order. 

Notice that after the first pass of a bubble sort, the biggest item is always moved to the end of the list. This means the next iteration does not need to check the final position of the list. 

In addition, you can keep track of how many items were swapped on each iteration. If this count is zero on any iteration, then the list must be fully sorted, in which case the algorithm can stop early.

#### Merge Sort
To understand merge sort, let's first understand the *merge* operation. The merge operation takes two *sorted* lists and combines them into another sorted list. Imagine we are merging `list1` and `list2`, both of which are sorted. Create two variables, one pointing to the first element (zero) of each list, call this `ptr1` and `ptr2`.

Then, go through each index of the merged list (e.g. `i in range(0, len(list1)+len(list2))`). Set position `i` of the merged list to the *smaller* value of `list1[ptr1]` and `list2[ptr2]`, then advance the corresponding pointer by one.

So initially you will compare the first element of each list, if the value in `list1` is smaller then copy this into the merged list, and then compare the second item of `list1` with the first item of `list2`.

Keep going until one pointer goes beyond its list – at this point, copy the rest of the other list into the merged list.

Have a look at an example of the merge operation below:

<br /><video controls loop autoplay width=600 src="./resources/merge.mp4">
</video>

Now, to understand *merge sort*, realise that any list of size 1 must be considered sorted. So, here is the process for the input list:
* If `in_list` is length 1, return `in_list`
* Otherwise, split `in_list` list into two halves: `left` and `right`
* Merge sort `left`
* Merge sort `right`
* Merge `left` and `right` into `merged`
* Return `merged`

For any size input list, this will be recursively split into two until both halves are length 1, and then the final sorted list will be rebuilt by merging sorted lists.

<!-- grade_id: cell-ffd4978cfbf4832e, points: 0, grade: false, solution: true -->
```python
def bubble_sort(in_list):
    # YOUR CODE HERE
    raise NotImplementedError()
    
def merge_sort(in_list):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e322586715be1ec0, points: 50, grade: true, solution: false -->
```python
# Note that your code will be checked to ensure you are implementing the correct algorithms!
assert(bubble_sort([37, 42, 9, 19, 35, 4, 53, 22]) == [4, 9, 19, 22, 35, 37, 42, 53])
assert(bubble_sort([5, 4, 3, 2, 1]) == [1, 2, 3, 4, 5])
assert(bubble_sort([]) == [])

assert(merge_sort([37, 42, 9, 19, 35, 4, 53, 22]) == [4, 9, 19, 22, 35, 37, 42, 53])
assert(merge_sort([5, 4, 3, 2, 1]) == [1, 2, 3, 4, 5])
assert(merge_sort([]) == [])
```

### Question 3
*(worth 25%)*

In the markdown cell below, in your own words, explain the time complexity, space complexity, and stability of bubble sort and merge sort – both as algorithms, and your specific implementations.

Of course, these are well known algorithms, so we are aware you can look up the answers. The marks in this section are awarded for for explaining *why* these properties hold, in your own words. I encourage you to try to work them out yourself before you look up the answers, this might make explaining it easier.

<!-- grade_id: cell-c50e99cdef82c2d1, points: 25, grade: true, solution: true -->
YOUR ANSWER HERE – double click to edit, write in [Markdown](https://jupyter-notebook.readthedocs.io/en/stable/examples/Notebook/Working%20With%20Markdown%20Cells.html), run to render

### Question 4
*(worth 25%)*

When we write simple arithmetic expressions, we use *infix notation*, the operator goes in the middle of the two arguments, such as `1 + 2`, or `5 - 3`.

An alternative is to use *prefix notation*; here the same expressions would be written `+ 1 2` and `- 5 3`. This is also called *polish notation*.

The reason this is useful is that it removes the need for parentheses in nested expressions. Suppose we want to multiply the two expressions. 

Using infix, we must use parentheses to ensure the operators are evaluated in the correct order: `(1 + 2) * (5 - 3)`.

Using prefix, we can write: `* + 1 2 - 5 3`. So long as all the operators are binary (two arguments), no parentheses are needed, no matter how complex the expression.

You can also build a *binary tree* to represent a nested expression unambiguously.

<img src="./resources/mathstree.png" width=400 />

Notice if you traverse the tree in-order (LNR) you get infix notation (without parentheses). If you traverse the tree pre-order (NLR) you get prefix notation. If you traverse the tree post-order (LRN) then you get *postfix notation*, which is also called *reverse polish notation*, which is also unambiguous under the same conditions.

In the cell below, write a function which takes a string containing a prefix notation mathematical expression, builds a binary tree, and then traverses it to produce an equivalent postfix notation expression.

The input expression will be valid. It will only contain positive integers and the operators `+`, `-`, `/`, and `*`. Terms will always be separated by a single space.

You are welcome to reuse the code from the unit material.

Hint: start simple, use the tests below to help you structure your thinking.

<!-- grade_id: cell-66f388f6c90d3698, points: 0, grade: false, solution: true -->
```python
def prefix_to_postfix(expression):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0a52d49715d5f7f1, points: 25, grade: true, solution: false -->
```python
assert(prefix_to_postfix("5") == "5")
assert(prefix_to_postfix("+ 10 20") == "10 20 +")
assert(prefix_to_postfix("* + 1 2 - 5 3") == "1 2 + 5 3 - *")
```