# Exercise Sheet 2
## Instructions
The instructions for this exercise sheet are the same as the previous, so this section is just a summary to serve as a reminder. Refer to the material on Engage and at the start of Exercise Sheet 1 for more detail, and if you are unsure ask on the Q&A forum.

This exercise sheet counts towards your overall grade for the unit.

Complete each question in the code cell provided, which will usually contain some skeleton code to help you start the question. Unless specified otherwise you may change the skeleton code as you wish, but your code must pass the formatting tests in the following code cell – if the cell runs without errors then your code is eligible for submission. However these will only test the basics of the question. Your code will be subject to additional hidden tests which will determine your grade. Make sure you thoroughly test your own code to check you think it works correctly before submission.

**Note**: For any question which accepts a mutable input (e.g. a list), your code must not modify that input unless specified otherwise in the question.

## Questions
### Question 1
Write a function that takes a list of items and returns a new list with the order of items reversed, without using any inbuilt Python features that will do this for you (e.g. `.reverse()` or `reversed(…)`). As specified in the instructions, your function must not modify the contents of the original list.

<!-- grade_id: cell-3fe55db4c6d2fc90, points: 0, grade: false, solution: true -->
```python
def reversed_list(in_list):
    # 結果を入れる空のリストを用意する
    new_list = [] # must not modify the contents of the original list?

    # インデックス用の変数を作る
    index = len(in_list) - 1

    # 一番後ろから0番目まで順番に取り出す
    while index >= 0:
        # 後ろから順に追加していく(.appendを使う)
        new_list.append(in_list[index])
        # 1つ前に戻る
        index = index - 1

    # できあがったリストを返す
    return new_list

## Reference: Yamada, Y. (2025). Dokushu Python (2nd ed.). Shoeisha.

```

<!-- grade_id: cell-4788d108c8c7a005, points: 10, grade: true, solution: false -->
```python
# ===== Question 1の追加テスト =====

# 公開テスト（既にある）
assert(reversed_list([1, 2]) == [2, 1])

# 追加テスト1: 空リスト
assert(reversed_list([]) == [])

# 追加テスト2: 要素1個
assert(reversed_list([5]) == [5])

# 追加テスト3: 長いリスト
assert(reversed_list([1, 2, 3, 4, 5]) == [5, 4, 3, 2, 1])

# 追加テスト4: 元のリストが変更されていないか（最重要）
original = [1, 2, 3]
result = reversed_list(original)
assert(original == [1, 2, 3])  # 元のリストは変わっていない
assert(result == [3, 2, 1])

# 追加テスト5: 文字列のリスト
assert(reversed_list(['a', 'b', 'c']) == ['c', 'b', 'a'])

# 追加テスト6: 同じ値が複数
assert(reversed_list([1, 1, 1]) == [1, 1, 1])

# 追加テスト7: 負の数
assert(reversed_list([-1, -2, -3]) == [-3, -2, -1])

# 追加テスト8: 混合型
assert(reversed_list([1, 'a', 2.5]) == [2.5, 'a', 1])

```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[18], line 4
      1 # ===== Question 1の追加テスト =====
      2 
      3 # 公開テスト（既にある）
----> 4 assert(reversed_list([1, 2]) == [2, 1])
      6 # 追加テスト1: 空リスト
      7 assert(reversed_list([]) == [])

AssertionError:
```

### Question 2
Now write a procedure that takes a list as an input and reverses its contents. So, the procedure should not return anything, and the inputted list *should* be modified. Of course, you still must not use inbuilt Python features which reverse lists.

<!-- grade_id: cell-26bdef0580e1c19c, points: 0, grade: false, solution: true -->
```python
def reverse_list(in_list):
    # 最初と最後を入れ替える、その次も入れ替える、という感じ
    i = 0
    j = len(in_list) - 1

    # iとjが交差するまで繰り返す
    while i < j:
        # 入れ替え
        temp = in_list[i] # 最も基本的な3点交換方式
        in_list[i] = in_list[j]
        in_list[j] = temp

        # incriment i and j
        i = i + 1
        j = j - 1

## Reference: Yamada, Y. (2025). Dokushu Python (2nd ed.). Shoeisha.

```

<!-- grade_id: cell-fde78ec384676dc5, points: 10, grade: true, solution: false -->
```python
# 公開テスト（既にある）
test_in = [1, 2]
reverse_list(test_in)
assert(test_in == [2, 1])

# 追加テスト1: 空リスト
test1 = []
reverse_list(test1)
assert(test1 == [])

# 追加テスト2: 要素1個
test2 = [5]
reverse_list(test2)
assert(test2 == [5])

# 追加テスト3: 長いリスト
test3 = [1, 2, 3, 4, 5]
reverse_list(test3)
assert(test3 == [5, 4, 3, 2, 1])

# 追加テスト4: 偶数個の要素
test4 = [1, 2, 3, 4]
reverse_list(test4)
assert(test4 == [4, 3, 2, 1])

# 追加テスト5: 奇数個の要素
test5 = [1, 2, 3]
reverse_list(test5)
assert(test5 == [3, 2, 1])

# 追加テスト6: 文字列のリスト
test6 = ['a', 'b', 'c']
reverse_list(test6)
assert(test6 == ['c', 'b', 'a'])

```

### Question 3
Given two lists `[a1, a2, a3, …]` and `[b1, b2, b3, …]` of the same length, return a list containing pairs of items as tuples by position: a tuple containing the first item from each list, a tuple containing the second item, and so on. In other words, return `[(a1, b1), (a2, b2), …]`.

This is quite a useful feature when you want to iterate over two collections at once, and it also has a inbuilt function called `zip`! Of course, you cannot use `zip` (or other similar functions) in your submission.

<!-- grade_id: cell-2b3c95462f105301, points: 0, grade: false, solution: true -->
```python
def merge(list1, list2):
    # 結果を入れるリスト
    list_as_result = []

    # 0番目、1番目、2番目...と順番に処理
    i = 0
    while i < len(list1): # 上限まで
        # (list1[i], list2[i]) というタプルを作って追加
        pair = (list1[i], list2[i])
        list_as_result.append(pair)
        i = i + 1

    return list_as_result

```

<!-- grade_id: cell-328705de8fb298ac, points: 10, grade: true, solution: false -->
```python
# 公開テスト
assert(merge([1, 2], [3, 4]) == [(1, 3), (2, 4)])

# 追加テスト1: 空リスト
assert(merge([], []) == [])

# 追加テスト2: 長いリスト
assert(merge([1, 2, 3], [4, 5, 6]) == [(1, 4), (2, 5), (3, 6)])

# 追加テスト3: 文字列
assert(merge(['a', 'b'], ['x', 'y']) == [('a', 'x'), ('b', 'y')])

# 追加テスト4: 要素1個
assert(merge([10], [20]) == [(10, 20)])

```

### Question 4
Write another function like the one from the previous question. It should take two *non-empty* lists and create a list of tuples. The two input lists do not need to be the same length. If they are not the same length, *wrap around* the shorter list, i.e. start from the beginning again. (Hint: don't forget the *modulo* operator.) As usual, do not use `zip` or other similar inbuilt functions which implement this functionality already.

So for example, given the inputs `[1, 2, 3, 4, 5]`, and `['a', 'b']`, the result should be: <br>
`[(1, 'a'), (2, 'b'), (3, 'a'), (4, 'b'), (5, 'a')]`.

<!-- grade_id: cell-4e38ae067a6ac9c4, points: 0, grade: false, solution: true -->
```python
def merge_wrap(list1, list2):
    # 結果を入れるリスト
    list_as_result = []

    # 長い方のリストの長さ分繰り返す
    longer_length = len(list1)
    if len(list2) > len(list1):
        longer_length = len(list2)

    # インデックスを1つずつ増やしていく
    i = 0
    while i < longer_length:
        # list1のインデックス（長さを超えたら最初に戻る）
        index1 = i % len(list1)
        # list2のインデックス（長さを超えたら最初に戻る）
        index2 = i % len(list2)

        # タプルを作って追加
        pair = (list1[index1], list2[index2])
        list_as_result.append(pair)

        i = i + 1

    return list_as_result

```

<!-- grade_id: cell-b3b03350120ec08f, points: 10, grade: true, solution: false -->
```python
# 公開テスト
assert(merge_wrap([1, 2], [3, 4, 5]) == [(1, 3), (2, 4), (1, 5)])

# 両方同じ長さ
assert(merge_wrap([1, 2], ['a', 'b']) == [(1, 'a'), (2, 'b')])

# list1の方が長い
assert(merge_wrap([1, 2, 3, 4], ['a', 'b']) == [(1, 'a'), (2, 'b'), (3, 'a'), (4, 'b')])

# list2の方が長い
assert(merge_wrap([1, 2], ['a', 'b', 'c', 'd']) == [(1, 'a'), (2, 'b'), (1, 'c'), (2, 'd')])

# 長さ1
assert(merge_wrap([10], ['x', 'y', 'z']) == [(10, 'x'), (10, 'y'), (10, 'z')])

```

### Question 5
Now one more variation: this time you must write a version of the merge and wrap function which takes *any number* of nonempty lists as input, using the *varargs* syntax you saw in the material about tuples. The return value should be a single list containing a number of tuples equal to the length of the longest list from all of the arguments, where each tuple has a number of elements equal to the number of arguments. Hopefully this is obvious, but you must not use `zip` or any other similar inbuilt functions which do the work for you.

See an example in the test cell below.

<!-- grade_id: cell-855bf4e1dfa15c7b, points: 0, grade: false, solution: true -->
```python
def merge_wrap_n(*args):

##################### validation if ######################
    # 引数が空の場合は空リストを返す
    if not args:
        return []

    # 全てのリストが空の場合も空リストを返す
    if all(len(lst) == 0 for lst in args):
        return []

    ######################################################

    # 一番長いリストの長さを見つける
    max_length = 0
    for lst in args:
        if len(lst) > max_length:
            max_length = len(lst)

    # 結果を入れるリスト
    list_as_result = []

    # 一番長いリストの長さ分だけ繰り返す
    for i in range(max_length):
        # このインデックスで作るタプルの要素を入れるリスト
        tuple_elements = []

        # 各リストから要素を取り出す
        for lst in args:
            # リストの長さで割った余りをインデックスにすることで、
            # リストの最初に戻って繰り返す（wrap around）
            index = i % len(lst)
            tuple_elements.append(lst[index])

        # リストをタプルに変換して結果に追加
        list_as_result.append(tuple(tuple_elements))

    return list_as_result

```

<!-- grade_id: cell-bb2620ea2d8898af, points: 10, grade: true, solution: false -->
```python
# テストケース1: 問題文の例
assert merge_wrap_n([1, 2], [3], [4, 5, 6]) == [(1, 3, 4), (2, 3, 5), (1, 3, 6)]

# テストケース2: 全てのリストが同じ長さ
assert merge_wrap_n([1, 2, 3], [4, 5, 6], [7, 8, 9]) == [(1, 4, 7), (2, 5, 8), (3, 6, 9)]

# テストケース3: リストが1つだけ
assert merge_wrap_n([1, 2, 3]) == [(1,), (2,), (3,)]

# テストケース4: 2つのリストで長さが異なる
assert merge_wrap_n([1, 2], [3, 4, 5, 6]) == [(1, 3), (2, 4), (1, 5), (2, 6)]

# テストケース5: 1要素のリストを含む
assert merge_wrap_n([1], [2, 3, 4]) == [(1, 2), (1, 3), (1, 4)]

# テストケース6: 4つのリスト
assert merge_wrap_n([1, 2], [3, 4, 5], [6], [7, 8, 9, 10]) == [
    (1, 3, 6, 7), (2, 4, 6, 8), (1, 5, 6, 9), (2, 3, 6, 10)
]

```

### Question 6
Write a function which takes a list of numbers of length $\ge 2$, and returns a tuple containing the *arithmetic mean* and the *unbiased sample variance* of those numbers.

i.e. given a list of items $[x_1, x_2, \dots, x_N]$, return a tuple $(m, s^2)$ where:

\begin{align} 
m &= \frac{1}{N} \sum_{i=1}^{N} x_i \\
s^2 &= \frac{1}{N-1} \sum_{i=1}^{N} (x_i - m)^2 .
\end{align}

The symbol $\sum$ means to add up all the values in this range for the variable specified. In other words, take the value where $i=1$, add the value where $i=2$, and so on until $i=N$. So the arithmetic mean can be explained in words as: add up all the values in the list, then divide by the total number of values ($N$). For the variance you need to subtract the mean from each value and square the result, add this up, then divide by $N-1$. You are welcome to use mathematically equivalent reformulations if you prefer (and it may be more efficient to do so), just ensure you match the outputs of these formulas exactly.

You must do the calculations manually, do not use inbuilt functions like `sum`, any `math` functions, etc.

<!-- grade_id: cell-bece9af02711a741, points: 0, grade: false, solution: true -->
```python
def mean_and_variance(numbers):
    """
    リストから算術平均と不偏標本分散を計算して返す

    Parameters:
    numbers: 長さ >= 2 の数値リスト

    Returns:
    tuple: (mean, variance)
    """
    # リストの長さを取得
    N = 0
    for _ in numbers:
        N += 1

    # 算術平均を計算: m = (1/N) * Σx_i
    total = 0
    for x in numbers:
        total += x
    m = total / N

    # 不偏標本分散を計算: s² = 1/(N-1) * Σ(x_i - m)²
    squared_diff_sum = 0
    for x in numbers:
        diff = x - m
        squared_diff_sum += diff * diff

    s2 = squared_diff_sum / (N - 1)

    return (m, s2)


# テスト
print(mean_and_variance([1, 2, 3, 4, 5]))
# 期待値: mean = 3.0, variance = 2.5

print(mean_and_variance([10, 20]))
# 期待値: mean = 15.0, variance = 50.0

print(mean_and_variance([2.5, 3.5, 4.5]))
# 期待値: mean = 3.5, variance = 1.0

```

Output:
```
(3.0, 2.5)
(15.0, 50.0)
(3.5, 1.0)
```

<!-- grade_id: cell-8b8d9bd745e70c03, points: 10, grade: true, solution: false -->
```python

```

### Question 7
Write a function which takes a string as an input and returns a dictionary which maps from each letter a-z (the keys), *ignoring case*, to the number of times this letter appears in the string (the values), *if the letter appears in the string* (all counts should be $\ge 1$). The keys to the dictionary must all be *lower case*. Ignore all punctuation or whitespace.

<!-- grade_id: cell-f38552901de3cd4a, points: 0, grade: false, solution: true -->
```python
def frequency_analysis(string):
    """
    文字列中の各アルファベット（a-z）の出現回数を辞書で返す

    Parameters:
    string: 入力文字列

    Returns:
    dict: {文字(小文字): 出現回数} ただし出現回数 >= 1 の文字のみ
    """
    # 結果を格納する辞書
    freq = {}

    # 文字列を1文字ずつ処理
    for char in string:
        # 小文字に変換
        lower_char = char.lower()

        # a-z の範囲内かチェック
        if 'a' <= lower_char <= 'z':
            # 既に辞書にあればカウントを増やす、なければ1を設定
            if lower_char in freq:
                freq[lower_char] += 1
            else:
                freq[lower_char] = 1

    return freq


# テスト
result = frequency_analysis("Hello world!")
print(result)
# 期待値: {'h': 1, 'e': 1, 'l': 3, 'o': 2, 'w': 1, 'r': 1, 'd': 1}

# 検証
assert frequency_analysis("Hello world!") == {'h': 1, 'e': 1, 'l': 3, 'o': 2, 'w': 1, 'r': 1, 'd': 1}
print("✓ テスト合格！")

```

Output:
```
{'h': 1, 'e': 1, 'l': 3, 'o': 2, 'w': 1, 'r': 1, 'd': 1}
✓ テスト合格！
```

<!-- grade_id: cell-3f5ae3bb369a1544, points: 10, grade: true, solution: false -->
```python
assert(frequency_analysis("Hello world!") == {'h': 1, 'e': 1, 'l': 3, 'o': 2, 'w': 1, 'r': 1, 'd': 1})
```

### Question 8
Write a function which takes a string as an input and returns a dictionary where the keys are the *words* from the string and the values are the *counts*, how often each word appears in the string. A word is a sequence of consecutive letters which is broken by a space, the end of the string, or any of the following punctuation symbols: `.,?!` (no other punctuation will be included in the input string). Ignore case when counting words: the dictionary's keys must be entirely in lower case. Ignore consecutive spaces or punctuation.

You are *encouraged* to make use of the various [string methods](https://docs.python.org/3/library/stdtypes.html#string-methods) available in Python.

<!-- grade_id: cell-441b97aa02b39592, points: 0, grade: false, solution: true -->
```python
def count_words(string):
    """
    文字列中の各単語の出現回数を辞書で返す

    Parameters:
    string: 入力文字列

    Returns:
    dict: {単語(小文字): 出現回数}

    Reference:
    Word frequency analysis implementation.
    See: https://www.bomberbot.com/python/finding-the-frequency-of-each-word-in-a-string-with-python-a-comprehensive-guide/

    """
    # 区切り文字
    punctuation = '.,?!'

    # 文字列を小文字に変換
    lower_string = string.lower()

    # 句読点をスペースに置換
    for p in punctuation:
        lower_string = lower_string.replace(p, ' ')

    # スペースで分割して単語リストを作成
    words = lower_string.split()

    # 単語をカウント
    word_count = {}
    for word in words:
        if word:  # 空文字列は無視
            if word in word_count:
                word_count[word] += 1
            else:
                word_count[word] = 1

    return word_count


# テスト
print(count_words("HeLlo worLd!"))
# 期待値: {"hello": 1, "world": 1}

print(count_words("The Bart, the"))
# 期待値: {"the": 2, "bart": 1}
```

Output:
```
{'hello': 1, 'world': 1}
{'the': 2, 'bart': 1}
```

<!-- grade_id: cell-1d02a1dcee1ca237, points: 10, grade: true, solution: false -->
```python
assert(count_words("Hello world!") == {"hello": 1, "world": 1})
assert(count_words("The Bart, the") == {"the": 2, "bart": 1})
```

### Question 9
In this question you will implement the [sieve of Eratosthenes](https://en.wikipedia.org/wiki/Sieve_of_Eratosthenes), a method for finding prime numbers. A prime number is any integer which is divisible by exactly two numbers: $1$ and itself. The basic idea is that you start with a list (or equivalent) of all the numbers from $2$ to $N$, then you repeat the following two steps:
* remove the first element from the list, this is your next prime number, call it $p$
* remove every element from the rest of the which is a multiple of $p$

By repeating these steps you fill find every prime number up to $N$. There are various optimisations you can perform, but this basic idea is sufficient for this exercise.

For this question, write a function which takes a single integer `n` and returns a *set* containing all of the prime numbers up to *and including* `n`.

<!-- grade_id: cell-657584a91cd8771d, points: 0, grade: false, solution: true -->
```python
def sieve(n):
    """
    エラトステネスの篩を使って、n以下の全素数を返す

    Parameters:
    n: 整数

    Returns:
    set: n以下の素数の集合

    Reference:
    Implementation based on the Sieve of Eratosthenes algorithm.
    See: https://www.geeksforgeeks.org/python/python-program-for-sieve-of-eratosthenes/

    """
    if n < 2:
        return set()

    # 2からnまでの数のリストを作成
    candidates = []
    for i in range(2, n + 1):
        candidates.append(i)

    primes = set()

    while candidates:
        # 最初の要素を取り出す（これが素数）
        p = candidates[0]
        primes.add(p)

        # pの倍数を全て削除
        new_candidates = []
        for num in candidates:
            if num % p != 0:
                new_candidates.append(num)

        candidates = new_candidates

    return primes


# テスト
print(sieve(10))
# 期待値: {2, 3, 5, 7}

print(sieve(20))
# 期待値: {2, 3, 5, 7, 11, 13, 17, 19}

assert sieve(10) == {2, 3, 5, 7}
print("✓ Question 9 テスト合格！")

```

Output:
```
{2, 3, 5, 7}
{2, 3, 5, 7, 11, 13, 17, 19}
✓ Question 9 テスト合格！
```

<!-- grade_id: cell-8203fc04b66a67e8, points: 10, grade: true, solution: false -->
```python
assert(sieve(10) == {2, 3, 5, 7})
```

### Question 10
In this question you need to write three functions, all of which return 2D lists with a given pattern made up of single-character strings containing either a hash symbol `'#'` or a space `' '`.

The three patterns are:
* `rectangle(n, m)`, where $n\ge2$ and $m\ge2$, must produce an `n` by `m` rectangle outline, e.g:
  ```
  rectangle(3, 3) returns
      [['#', '#', '#'], 
       ['#', ' ', '#'], 
       ['#', '#', '#']]
       
  rectangle(4, 5) returns
      [['#', '#', '#', '#', '#'],
       ['#', ' ', ' ', ' ', '#'],
       ['#', ' ', ' ', ' ', '#'],
       ['#', '#', '#', '#', '#']]
  ```
* `diamond(n)`, where $n\ge3$, must produce an `n` by `n` filled diamond, where odd sized diamonds are one hash wide at their points, and even shaped diamonds are two hashes wide, e.g.:
  ```
  diamond(5) returns
      [[' ', ' ', '#', ' ', ' '],
       [' ', '#', '#', '#', ' '],
       ['#', '#', '#', '#', '#'],
       [' ', '#', '#', '#', ' '],
       [' ', ' ', '#', ' ', ' ']]
       
  diamond(6) returns
      [[' ', ' ', '#', '#', ' ', ' '],
       [' ', '#', '#', '#', '#', ' '],
       ['#', '#', '#', '#', '#', '#'],
       ['#', '#', '#', '#', '#', '#'],
       [' ', '#', '#', '#', '#', ' '],
       [' ', ' ', '#', '#', ' ', ' ']]
  ```
* `chequerboard(n, m, c_h, c_w)`, all arguments $\ge1$, returns an `n` by `m` chequerboard where each chequer (region) is of size `c_h` by `c_w`, the top left cell is *always* filled, and *only* the right hand and bottom edge may contain incomplete chequers, e.g.:
  ```
  chequerboard(6, 8, 2, 2) returns
      [['#', '#', ' ', ' ', '#', '#', ' ', ' '],
       ['#', '#', ' ', ' ', '#', '#', ' ', ' '],
       [' ', ' ', '#', '#', ' ', ' ', '#', '#'],
       [' ', ' ', '#', '#', ' ', ' ', '#', '#'],
       ['#', '#', ' ', ' ', '#', '#', ' ', ' '],
       ['#', '#', ' ', ' ', '#', '#', ' ', ' ']]
      
  chequerboard(5, 7, 2, 3) returns
      [['#', '#', '#', ' ', ' ', ' ', '#'],
       ['#', '#', '#', ' ', ' ', ' ', '#'],
       [' ', ' ', ' ', '#', '#', '#', ' '],
       [' ', ' ', ' ', '#', '#', '#', ' '],
       ['#', '#', '#', ' ', ' ', ' ', '#']]
  ```
  
  
For this question, `rectangle(n, m)` is worth a maximum of 2 marks, `diamond(n)` is worth 3 marks, and `chequerboard(…)` is worth 5 marks.

<!-- grade_id: cell-1d1b2cff91c4b309, points: 0, grade: false, solution: true -->
```python
def rectangle(n, m):
    """
    n×mの長方形の枠線を描く
    """
    result = []
    for i in range(n):
        row = []
        for j in range(m):
            # 最初の行、最後の行、または最初の列、最後の列なら#
            if i == 0 or i == n - 1 or j == 0 or j == m - 1:
                row.append('#')
            else:
                row.append(' ')
        result.append(row)
    return result


def diamond(n):
    """
    n×nのダイヤモンド形を描く
    奇数サイズ：先端は1つの#
    偶数サイズ：先端は2つの#
    """
    result = []
    mid = n // 2

    for i in range(n):
        row = []
        for j in range(n):
            # 奇数サイズの場合
            if n % 2 == 1:
                # ダイヤモンドの内側かどうかを判定
                if i <= mid:
                    # 上半分
                    if mid - i <= j <= mid + i:
                        row.append('#')
                    else:
                        row.append(' ')
                else:
                    # 下半分
                    if i - mid <= j <= n - 1 - (i - mid):
                        row.append('#')
                    else:
                        row.append(' ')
            else:
                # 偶数サイズの場合
                if i < mid:
                    # 上半分
                    if mid - 1 - i <= j <= mid + i:
                        row.append('#')
                    else:
                        row.append(' ')
                else:
                    # 下半分
                    if i - mid <= j <= n - 1 - (i - mid):
                        row.append('#')
                    else:
                        row.append(' ')
        result.append(row)
    return result


def chequerboard(n, m, c_h, c_w):
    """
    n×mのチェッカーボード（市松模様）を描く
    各マス目のサイズはc_h×c_w
    左上は常に#で始まる
    """
    result = []
    for i in range(n):
        row = []
        for j in range(m):
            # どのチェッカーのブロックに属するか計算
            block_row = i // c_h
            block_col = j // c_w

            # ブロックの位置が偶数+偶数 または 奇数+奇数 なら#
            if (block_row + block_col) % 2 == 0:
                row.append('#')
            else:
                row.append(' ')
        result.append(row)
    return result


# テスト
rect = [['#', '#', '#'],
        ['#', ' ', '#'],
        ['#', '#', '#']]

diam = [[' ', '#', ' '],
        ['#', '#', '#'],
        [' ', '#', ' ']]

cheq = [['#', ' ', '#'],
        [' ', '#', ' '],
        ['#', ' ', '#']]

assert(rectangle(3, 3) == rect)
assert(diamond(3) == diam)
assert(chequerboard(3, 3, 1, 1) == cheq)

print("✓ 全テスト合格！")

# 追加テスト
print("\nrectangle(4, 5):")
for row in rectangle(4, 5):
    print(''.join(row))

print("\ndiamond(6):")
for row in diamond(6):
    print(''.join(row))

print("\nchequerboard(5, 7, 2, 3):")
for row in chequerboard(5, 7, 2, 3):
    print(''.join(row))

```

Output:
```
✓ 全テスト合格！

rectangle(4, 5):
#####
#   #
#   #
#####

diamond(6):
  ##  
 #### 
######
######
 #### 
  ##  

chequerboard(5, 7, 2, 3):
###   #
###   #
   ### 
   ### 
###   #
```

<!-- grade_id: cell-699d6450bbd8bbd9, points: 10, grade: true, solution: false -->
```python
rect = [['#', '#', '#'], 
        ['#', ' ', '#'], 
        ['#', '#', '#']]

diam = [[' ', '#', ' '], 
        ['#', '#', '#'], 
        [' ', '#', ' ']]

cheq = [['#', ' ', '#'], 
        [' ', '#', ' '], 
        ['#', ' ', '#']]

assert(rectangle(3, 3) == rect)
assert(diamond(3) == diam)
assert(chequerboard(3, 3, 1, 1) == cheq)
```