# Exercise Sheet 1
## Instructions
This exercise sheet counts towards your overall grade for the unit. Please refer to Engage for the most up-to-date administrative details on grading and deadlines.

Each exercise sheet will be made up of several programming challenges in Python. The content of the exercises will primarily focus on the material from the relevant week, but concepts from previous weeks will also be required, and you may even need to pick up new concepts that are introduced within the exercise itself – this is a normal part of programming.

Each exercise is spread across three Jupyter notebook cells:
* A text cell (like this one) introducing the exercise
* A code cell where you can fill your solution, which may include some code already pre-filled for you to modify
* A code cell containing tests

Each test cell will normally include visible tests that you must run yourself to check your submission is working and in the format we expect. If you can run the cell and *nothing* happens, the tests are passing. These tests are mandatory, and do not typically include challenging test cases. If your submission does not pass these tests then it may receive zero marks for the exercise. 

The test cell also contains hidden tests, and these will be used to actually determine your grade for the exercise. So you **must** also test your code on additional inputs that you determine yourself. This is part of writing code, you need to check that it works. 

Unless otherwise specified, you *do not* need to ensure that your code is robust for unusual inputs. If the exercise states that the input will be a number, you can expect the input to be a number, you do not need to check if it is a string.

You can test your code in the middle code cell or add additional cells to the notebook, but *please delete these before your final submission*, which should just be the exercise code itself.

In general, unless the exercise instructions state otherwise, you can feel free to modify any skeleton code that has been provided *provided* the test cell still runs successfully.

The provided code *raise NotImplementedError()* is just a placeholder, please delete it after you add your own code.

## Questions
### Question 1
Write a function that takes a single number, adds 3 to that number, then multiplies the result by 2, then subtracts 4, then subtracts twice the original number, then adds 3, and returns the result.

<!-- grade_id: cell-8ec771280bdacf20, points: 0, grade: false, solution: true -->
```python
def process_number(x):
    originalX = x            # save original X

    step1 = originalX + 3    # 1. add 3
    step2 = step1 * 2       # 2.2倍する
    step3 = step2 - 4       # 3.4を引く
    step4 = step3 - (2 * originalX)  # 4.元の数の2倍を引く
    step5 = step4 + 3       # 5.3を足す

    return step5

## Reference: Yamada, Y. (2025). Dokushu Python (2nd ed.). Shoeisha.

```

<!-- grade_id: cell-37523cf9c5e087d0, points: 10, grade: true, solution: false -->
```python
# Question 1用の追加テスト
assert(process_number(0) == 5)
assert(process_number(-10) == 5)
assert(process_number(1000) == 5)
assert(process_number(3.5) == 5)

```

### Question 2
Write a function which converts a length given in centimetres to feet. You will likely need to look up the formula for this conversion online, try to be as exact as possible!

<!-- grade_id: cell-62ec441834039e8b, points: 0, grade: false, solution: true -->
```python
def centimetres_to_feet(cm):
    # 1 foot = 30.48 cm
    result = cm / 30.48
    return result

## Reference: Yamada, Y. (2025). Dokushu Python (2nd ed.). Shoeisha.

```

<!-- grade_id: cell-8d322b244bce8acb, points: 10, grade: true, solution: false -->
```python
# 公開テスト
assert(centimetres_to_feet(0) == 0.0)

# 30.48cmは1 foot
assert(centimetres_to_feet(30.48) == 1.0)

# 60.96cmは2 feet
assert(centimetres_to_feet(60.96) == 2.0)

# 負の数
assert(centimetres_to_feet(-30.48) == -1.0)

```

### Question 3
Write a function that takes two integers `x` and `y`. It should return a string containing the integer `x` repeated `y` times. See the provided tests below for examples.

<!-- grade_id: cell-a80faf28af2fff25, points: 0, grade: false, solution: true -->
```python
def repeat_int_str(x, y):
    # 数字を文字列にする
    s = str(x)
    # y回繰り返す
    result = s * y
    return result

## Reference: Yamada, Y. (2025). Dokushu Python (2nd ed.). Shoeisha.

```

<!-- grade_id: cell-abad8e6be2aa33b3, points: 10, grade: true, solution: false -->
```python
# 公開テスト
assert(repeat_int_str(0, 4) == "0000")
assert(repeat_int_str(-5, 3) == "-5-5-5")

# 0回繰り返し
assert(repeat_int_str(5, 0) == "")

# 1回だけ
assert(repeat_int_str(123, 1) == "123")

# 負の数
assert(repeat_int_str(-123, 2) == "-123-123")

# 1桁の数
assert(repeat_int_str(7, 5) == "77777")

```

### Question 4
Now write a function similar to the one above, that takes two integers `x` and `y`, but this time produces a string *of length* `y` made up of repeated copies of `x`.

<!-- grade_id: cell-9ce26ef0a58e99bb, points: 0, grade: false, solution: true -->
```python
def repeat_int_length(x, y):
    # 数字を文字列にする
    s = str(x)

    # yが0以下なら空文字列
    if y <= 0:
        return ""

    # sを何回か繰り返して、長さがy以上になるようにする
    result = ""
    while len(result) < y:
        result = result + s

    # 先頭からy文字だけ取り出す
    return result[:y]

## Reference: Yamada, Y. (2025). Dokushu Python (2nd ed.). Shoeisha.

```

<!-- grade_id: cell-b0c00ceb2ed4ab1b, points: 10, grade: true, solution: false -->
```python
# 公開テスト
assert(repeat_int_length(10, 4) == "1010")
assert(repeat_int_length(-5, 3) == "-5-")

# 長さ0
assert(repeat_int_length(123, 0) == "")

# 長さ1
assert(repeat_int_length(123, 1) == "1")

# ちょうど同じ長さ
assert(repeat_int_length(123, 3) == "123")

# 1桁の数
assert(repeat_int_length(7, 5) == "77777")

# 負の数
assert(repeat_int_length(-99, 7) == "-99-99-")

```

### Question 5
One more time, write another function which takes arguments `x` and `y`, and produces a string of length `y` made of repeated copies of `x`, ***but*** this time, if `x` is negative do not repeat the negative sign, only the digits. Again, see the examples in the tests below for clarification.

<!-- grade_id: cell-ece9e5b5e185e772, points: 0, grade: false, solution: true -->
```python
def repeat_int_no_minus(x, y):
    # yが0以下なら空文字列
    if y <= 0:
        return ""

    # xが負の数かチェック
    if x < 0:
        sign = "-"
        # 絶対値を文字列にする
        digits = str(-x)
    else:
        sign = ""
        digits = str(x)

    # digitsを繰り返してy文字以上にする
    result = ""
    while len(result) < y:
        result = result + digits

    # 先頭からy文字取り出す
    body = result[:y]

    # 符号と本体をつなげて返す
    return sign + body

## Reference: Yamada, Y. (2025). Dokushu Python (2nd ed.). Shoeisha.


```

<!-- grade_id: cell-7ba13193fa2aeee2, points: 10, grade: true, solution: false -->
```python
# 公開テスト
assert(repeat_int_no_minus(-5, 5) == "-5555")
assert(repeat_int_no_minus(-343, 3) == "-34")

# 正の数
assert(repeat_int_no_minus(123, 5) == "12312")

# 0
assert(repeat_int_no_minus(0, 4) == "0000")

# 長さ0
assert(repeat_int_no_minus(-5, 0) == "")

# 符号のみ
assert(repeat_int_no_minus(-99, 1) == "-")

# 負の1桁
assert(repeat_int_no_minus(-7, 6) == "-77777")

```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[14], line 2
      1 # 公開テスト
----> 2 assert(repeat_int_no_minus(-5, 5) == "-5555")
      3 assert(repeat_int_no_minus(-343, 3) == "-34")
      5 # 正の数

AssertionError:
```

### Question 6
Remember that strings can be compared using `<` and `>`. If you are comparing characters that are the same case (both UPPER CASE or both lower case), then `<` and `>` will compare the characters alphabetically. So `"a" < "z"` evaluates to `True`. 

Write a function that takes an input string all in lower case and finds the smallest (alphabetically first) *letter* from (a-z). Importantly, note that some special characters (like spaces) might be considered “smaller” and so must be ignored, i.e. `" " < "a"` evaluates to `True`. If the input string contains no lower case letters from a-z, return an empty string.

<!-- grade_id: cell-62b319d167ceab6f, points: 0, grade: false, solution: true -->
```python
def smallest_letter(string):
    # 今まで見つけた一番小さい文字
    smallest = None

    # 文字を1つずつ見ていく
    for ch in string:
        # a〜zの小文字だけ処理する
        if ch >= 'a' and ch <= 'z':
            # 最初の小文字、または今の文字の方が小さい場合
            if smallest == None or ch < smallest:
                smallest = ch

    # 小文字が見つからなかったら空文字列
    if smallest == None:
        return ""
    else:
        return smallest

## Reference: Yamada, Y. (2025). Dokushu Python (2nd ed.). Shoeisha.


```

<!-- grade_id: cell-48d46b3c2804273d, points: 10, grade: true, solution: false -->
```python
# 公開テスト
assert(smallest_letter("hello world") == "d")
assert(smallest_letter(" ") == "")

# 最初が最小
assert(smallest_letter("abc") == "a")

# 最後が最小
assert(smallest_letter("zyx") == "x")

# 大文字混在
assert(smallest_letter("HeLLo WoRLd") == "d")

# 数字混在
assert(smallest_letter("123abc!@#xyz") == "a")

# 小文字なし
assert(smallest_letter("HELLO") == "")

# 空文字列
assert(smallest_letter("") == "")

```

### Question 7
The reason why we can compare strings alphabetically using `<` and `>`, but only for characters of the same case, is because it is *actually* comparing the [character values](https://en.wikipedia.org/wiki/List_of_Unicode_characters#Latin_script). Python uses Unicode for its characters, but you might also hear people refer to ASCII character values too, since Unicode is backwards compatible and uses the same values for basic Latin script, numerals, and the symbols you find on a UK keyboard.

You can convert a single-character string into its character value using the inbuilt function `ord`, and you can convert back from a number into a single-character string using `chr`. So `ord("A") == 65` and `chr(65) == "A"`. Since each character's value is sequential (A is 65, B is 66, …), if you think about it, you can use these functions to convert between cases. You can also use string methods like `.lower()` and `.upper()`.

Now, armed with this new information, write a function which works the same as the previous question except it produces the smallest alphabetical letter *of either case*. If the string contains an upper and lower case character which is alphabetically first, return the value of the character which occurs first in the string. The case of the return value must match the input. See the tests for clarifying examples.

<!-- grade_id: cell-b1729d6895302bf6, points: 0, grade: false, solution: true -->
```python
def smallest_ignore_case(string):
    # 結果として返す文字
    best_char = ""
    # 比較用の小文字版（最小値）
    best_lower = None

    # 文字を1つずつ見る
    for ch in string:
        # 英字かどうかチェック
        if (ch >= 'a' and ch <= 'z') or (ch >= 'A' and ch <= 'Z'):
            # 小文字にして比較する
            low = ch.lower()

            # 最初の英字、または今の文字の方が小さい
            if best_lower == None or low < best_lower:
                best_lower = low
                best_char = ch

    return best_char

## Reference: Yamada, Y. (2025). Dokushu Python (2nd ed.). Shoeisha.

```

<!-- grade_id: cell-96eff72adde04801, points: 10, grade: true, solution: false -->
```python
# 公開テスト
assert(smallest_ignore_case("Hello World") == "d")
assert(smallest_ignore_case("HeLLo WoRLd") == "D")
assert(smallest_ignore_case("What an AARDVARK!") == "a")

# 全て大文字
assert(smallest_ignore_case("ZEBRA") == "A")

# 全て小文字
assert(smallest_ignore_case("zebra") == "a")

# 大文字・小文字混在
assert(smallest_ignore_case("BbAa") == "A")

# 数字混在
assert(smallest_ignore_case("123Z456a789") == "a")

# 英字なし
assert(smallest_ignore_case("12345!@#") == "")

# 空文字列
assert(smallest_ignore_case("") == "")

```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[18], line 3
      1 # 公開テスト
      2 assert(smallest_ignore_case("Hello World") == "d")
----> 3 assert(smallest_ignore_case("HeLLo WoRLd") == "D")
      4 assert(smallest_ignore_case("What an AARDVARK!") == "a")
      6 # 全て大文字

AssertionError:
```

### Question 8
Write a function which takes a string as input and returns the biggest distance between any two letters (a-z ignoring case). Distance in this case means the number of letters apart in the alphabet, so the distance between `"a"` and `"b"` is `1`, the distance between `"a"` and `"Z"` is `25`. If the string contains no letters return `-1`.

<!-- grade_id: cell-cdb03a89c16be097, points: 0, grade: false, solution: true -->
```python
def biggest_letter_distance(string):
    # 英字だけを集める
    letters = []
    for char in string:
        if (char >= 'a' and char <= 'z') or (char >= 'A' and char <= 'Z'):
            # 小文字にして追加
            letters.append(char.lower())

    # 英字がなければ-1
    if len(letters) == 0:
        return -1

    # 各文字の位置を計算する (a=0, b=1, ...)
    positions = []
    for char in letters:
        pos = ord(char) - ord('a')
        positions.append(pos)

    # 最大と最小の差を求める
    max_pos = max(positions)
    min_pos = min(positions)
    distance = max_pos - min_pos

    return distance

## Reference: Yamada, Y. (2025). Dokushu Python (2nd ed.). Shoeisha.



```

<!-- grade_id: cell-4e3bc1006f3cde9a, points: 10, grade: true, solution: false -->
```python
# 公開テスト
assert(biggest_letter_distance("Hello World") == 19)
assert(biggest_letter_distance("!!") == -1)

# a と z（最大距離）
assert(biggest_letter_distance("az") == 25)

# a と b（最小距離）
assert(biggest_letter_distance("ab") == 1)

# 同じ文字のみ
assert(biggest_letter_distance("aaa") == 0)

# 大文字と小文字混在
assert(biggest_letter_distance("AZ") == 25)

# 英字なし
assert(biggest_letter_distance("") == -1)

```

### Question 9
In the following cell the programmer has been experimenting with writing a function that can take either two or three arguments. However, they have made a mistake with their syntax. Find the mistake and fix the code so that it behaves the way they are expecting.

<!-- grade_id: cell-8b1227a02347648d, points: 0, grade: false, solution: true -->
```python
def my_function(a, b, c=None):
    # cがない場合（2つの引数）
    if c == None:
        if a < b:
            result = a * a + b
            return result
        else:
            return b
    # cがある場合（3つの引数）
    else:
        # cが一番小さい場合
        if c < a and c < b:
            result = c * c * c + a * a + b
            return result
        else:
            # cが一番小さくない場合
            if a < b:
                result = a * a + b
                return result
            else:
                return b

## Reference: Yamada, Y. (2025). Dokushu Python (2nd ed.). Shoeisha.

```

<!-- grade_id: cell-cf0fd922b0a9b735, points: 10, grade: true, solution: false -->
```python
# 公開テスト
assert(my_function(4, 5, 3) == 48)
assert(my_function(2, 1) == 1)

# 2引数、a < b
assert(my_function(2, 5) == 9)

# 2引数、a >= b
assert(my_function(5, 2) == 2)

# 3引数、cが最小
assert(my_function(5, 6, 2) == 39)

# 3引数、cが最小でない
assert(my_function(2, 5, 7) == 9)

```

### Question 10
In the following code cell, a programmer has tried to re-implement the `*` and `//` operators as recursive functions that only use addition and subtraction; no other arithmetic operators (obviously `*` and `//` are not allowed) and no functions like `abs`. We saw an example of multiplication using addition in the unit material already using a loop, and the recursive version is a similar idea. Both functions should work for all reasonably-sized integer inputs `x` and `y`, the programmer is not trying to make them efficient. They should produce *exactly* the same results as the `*` and `//` operators.

Unfortunately the programmer has made some errors in their code. For this exercise, debug the code to find and fix all of the errors. Bear in mind that just getting the code to run does not mean you are done, as usual you will need to test the code with various inputs yourself before submitting. This will likely require tracing the code by hand keeping track of the values, even through layers of recursion!

<!-- grade_id: cell-09b65380a7494ba7, points: 0, grade: false, solution: true -->
```python
def multiply(x, y):
    if y == 0:
        return 0
    elif y > 0:
        return x + multiply(x, y - 1)
    else:
        # yが負の場合
        return 0 - multiply(x, 0 - y)

def divide(x, y):
    # 両方負
    if x < 0 and y < 0:
        return divide(0 - x, 0 - y)
    # xだけ負
    elif x < 0:
        return 0 - divide(0 - x, y)
    # yだけ負
    elif y < 0:
        return 0 - divide(x, 0 - y)
    # 両方正
    else:
        return divide_helper(x, y, 0)

def divide_helper(x, y, temp):
    if x < y:
        return temp
    else:
        return divide_helper(x - y, y, temp + 1)

## Reference: Yamada, Y. (2025). Dokushu Python (2nd ed.). Shoeisha.

```

<!-- grade_id: cell-e1ddc5f3ee96a7d5, points: 10, grade: true, solution: false -->
```python
# 公開テスト
assert(multiply(3, 5) == 3 * 5)
assert(divide(13, 3) == 13 // 3)

# multiply のテスト
assert(multiply(0, 5) == 0)
assert(multiply(5, 0) == 0)
assert(multiply(3, -4) == -12)
assert(multiply(-3, -4) == 12)

# divide のテスト
assert(divide(10, 3) == 3)
assert(divide(-10, 3) == -3)
assert(divide(10, -3) == -3)
assert(divide(0, 5) == 0)

```