# Exercise Sheet 1
## Instructions
This exercise sheet counts towards your overall grade for the unit. Please refer to Engage for the most up-to-date administrative details on grading and deadlines.

Each exercise sheet will be made up of several programming challenges in Python. The content of the exercises will primarily focus on the material from the relevant week, but concepts from previous weeks will also be required, and you may even need to pick up new concepts that are introduced within the exercise itself – this is a normal part of programming.

Each exercise is spread across three Jupyter notebook cells:
* A text cell (like this one) introducing the exercise
* A code cell where you can fill your solution, which may include some code already pre-filled for you to modify
* A code cell containing tests

Each test cell will normally include visible tests that you must run yourself to check your submission is working and in the format we expect. If you can run the cell and *nothing* happens, the tests are passing. These tests are mandatory, and do not typically include challenging test cases. If your submission does not pass these tests then it may receive zero marks for the exercise. 

The test cell also contains hidden tests, and these will be used to actually determine your grade for the exercise. So you **must** also test your code on additional inputs that you determine yourself. This is part of writing code, you need to check that it works. 

Unless otherwise specified, you *do not* need to ensure that your code is robust for unusual inputs. If the exercise states that the input will be a number, you can expect the input to be a number, you do not need to check if it is a string.

You can test your code in the middle code cell or add additional cells to the notebook, but *please delete these before your final submission*, which should just be the exercise code itself.

In general, unless the exercise instructions state otherwise, you can feel free to modify any skeleton code that has been provided *provided* the test cell still runs successfully.

The provided code *raise NotImplementedError()* is just a placeholder, please delete it after you add your own code.

## Questions
### Question 1
Write a function that takes a single number, adds 3 to that number, then multiplies the result by 2, then subtracts 4, then subtracts twice the original number, then adds 3, and returns the result.

<!-- grade_id: cell-8ec771280bdacf20, points: 0, grade: false, solution: true -->
```python
def process_number(x):
    originalX = x            # 元の数を保存

    step1 = originalX + 3    # 1.3を足す
    step2 = step1 * 2       # 2.2倍する
    step3 = step2 - 4       # 3.4を引く
    step4 = step3 - (2 * originalX)  # 4.元の数の2倍を引く
    step5 = step4 + 3       # 5.3を足す

    return step5

```

<!-- grade_id: cell-37523cf9c5e087d0, points: 10, grade: true, solution: false -->
```python
# Question 1用の追加テスト
assert(process_number(0) == 5)
assert(process_number(-10) == 5)
assert(process_number(1000) == 5)
assert(process_number(3.5) == 5)

```

Output:
```
✓ 全テスト通過
```

### Question 2
Write a function which converts a length given in centimetres to feet. You will likely need to look up the formula for this conversion online, try to be as exact as possible!

<!-- grade_id: cell-62ec441834039e8b, points: 0, grade: false, solution: true -->
```python
def centimetres_to_feet(cm):
    # 1 foot = 30.48 cm (exact), so feet = cm / 30.48
    feet = cm / 30.48
    return feet
```

<!-- grade_id: cell-8d322b244bce8acb, points: 10, grade: true, solution: false -->
```python
# ===== Question 2の追加テスト =====

# 公開テスト（既にある）
assert(centimetres_to_feet(0) == 0.0)

# 追加テスト1: 1 foot分（30.48cm）
assert(centimetres_to_feet(30.48) == 1.0)

# 追加テスト2: 2 feet分
assert(centimetres_to_feet(60.96) == 2.0)

# 追加テスト3: 小数
result = centimetres_to_feet(15.24)
assert(abs(result - 0.5) < 0.0001)  # 浮動小数点誤差対策

# 追加テスト4: 大きい数
result = centimetres_to_feet(304.8)
assert(abs(result - 10.0) < 0.0001)

# 追加テスト5: 負の数（念のため）
assert(centimetres_to_feet(-30.48) == -1.0)

```

Output:
```
✓ 30.48cm = 1 foot
✓ 60.96cm = 2 feet
✓ 15.24cm = 0.5 feet
✓ 304.8cm = 10 feet
✓ 負の数も対応

全テスト通過！
```

### Question 3
Write a function that takes two integers `x` and `y`. It should return a string containing the integer `x` repeated `y` times. See the provided tests below for examples.

<!-- grade_id: cell-a80faf28af2fff25, points: 0, grade: false, solution: true -->
```python
def repeat_int_str(x, y):
    # x を文字列に変換する（例: 12 -> "12"）
    x_str = str(x)

    # 文字列は * を使うと「回数分くり返して連結」できる（例: "12" * 3 -> "121212"）
    result = x_str * y

    # できあがった文字列を返す
    return result

```

<!-- grade_id: cell-abad8e6be2aa33b3, points: 10, grade: true, solution: false -->
```python
# ===== Question 3の追加テスト =====

# 公開テスト（既にある）
assert(repeat_int_str(0, 4) == "0000")
assert(repeat_int_str(-5, 3) == "-5-5-5")

# 追加テスト1: y=0（繰り返しなし）
assert(repeat_int_str(5, 0) == "")

# 追加テスト2: y=1（1回だけ）
assert(repeat_int_str(123, 1) == "123")

# 追加テスト3: 大きい数
assert(repeat_int_str(999, 2) == "999999")

# 追加テスト4: 負の数（より大きい）
assert(repeat_int_str(-123, 2) == "-123-123")

# 追加テスト5: 1桁の数
assert(repeat_int_str(7, 5) == "77777")

# 追加テスト6: 多数回繰り返し
result = repeat_int_str(1, 10)
assert(result == "1111111111")
assert(len(result) == 10)

```

Output:
```
✓ 0回繰り返し = 空文字列
✓ 1回繰り返し
✓ 大きい数
✓ 負の大きい数
✓ 1桁の数
✓ 10回繰り返し

全テスト通過！
```

### Question 4
Now write a function similar to the one above, that takes two integers `x` and `y`, but this time produces a string *of length* `y` made up of repeated copies of `x`.

<!-- grade_id: cell-9ce26ef0a58e99bb, points: 0, grade: false, solution: true -->
```python
def repeat_int_length(x, y):
    # x を文字列にする（例: 12 -> "12"）
    s = str(x)

    # 長さ y が 0 以下なら、長さ 0 の文字列を返す
    if y <= 0:
        return ""

    # まず「十分な長さ」になるまで s を繰り返してつなげる
    # 例: s="12", y=5 なら "12" を何回か繰り返して "121212..." の形にする
    repeats = (y + len(s) - 1) // len(s)   # 必要回数（切り上げ）
    long_str = s * repeats

    # 先頭から y 文字だけ切り出して、長さがちょうど y の文字列にする
    return long_str[:y]

```

<!-- grade_id: cell-b0c00ceb2ed4ab1b, points: 10, grade: true, solution: false -->
```python
# ===== Question 4の追加テスト =====

# 公開テスト（既にある）
assert(repeat_int_length(10, 4) == "1010")
assert(repeat_int_length(-5, 3) == "-5-")

# 追加テスト1: y=0（長さ0）
assert(repeat_int_length(123, 0) == "")

# 追加テスト2: y=1（1文字だけ）
assert(repeat_int_length(123, 1) == "1")

# 追加テスト3: xの長さとyが同じ
assert(repeat_int_length(123, 3) == "123")

# 追加テスト4: yがxの長さより短い
assert(repeat_int_length(12345, 3) == "123")

# 追加テスト5: yがxの長さの倍数
assert(repeat_int_length(12, 6) == "121212")

# 追加テスト6: 1桁の数
assert(repeat_int_length(7, 5) == "77777")

# 追加テスト7: 負の数（長い）
assert(repeat_int_length(-99, 7) == "-99-99-")

# 追加テスト8: y が負の数
assert(repeat_int_length(5, -3) == "")

# 追加テスト9: 0を繰り返す
assert(repeat_int_length(0, 5) == "00000")
```

Output:
```
✓ 長さ0
✓ 長さ1
✓ 長さ = xの桁数
✓ 長さ < xの桁数
✓ 長さ = xの桁数の倍数
✓ 1桁の数
✓ 負の数、長さ7
✓ 負の長さ = 空文字列
✓ 0を繰り返す

全テスト通過！
```

### Question 5
One more time, write another function which takes arguments `x` and `y`, and produces a string of length `y` made of repeated copies of `x`, ***but*** this time, if `x` is negative do not repeat the negative sign, only the digits. Again, see the examples in the tests below for clarification.

<!-- grade_id: cell-ece9e5b5e185e772, points: 0, grade: false, solution: true -->
```python
def repeat_int_no_minus(x, y):
    # 長さ y が 0 以下なら空文字列
    if y <= 0:
        return ""

    # x が負なら符号は "-"、負でないなら符号は空
    sign = "-" if x < 0 else ""

    # 繰り返すのは数字部分だけ（絶対値を文字列にする）
    digits = str(abs(x))

    # digits を十分な長さになるまで繰り返して、先頭 y 文字を取る
    repeats = (y + len(digits) - 1) // len(digits)  # 必要回数（切り上げ）
    body = (digits * repeats)[:y]

    # 符号 + 本体 を返す（負の場合でも "-" は1回だけ）
    return sign + body

```

<!-- grade_id: cell-7ba13193fa2aeee2, points: 10, grade: true, solution: false -->
```python
# ===== Question 5の追加テスト =====

# 公開テスト（既にある）
assert(repeat_int_no_minus(-5, 5) == "-5555")
assert(repeat_int_no_minus(-343, 3) == "-34")

# 追加テスト1: 正の数（符号なし）
assert(repeat_int_no_minus(123, 5) == "12312")

# 追加テスト2: 0
assert(repeat_int_no_minus(0, 4) == "0000")

# 追加テスト3: y=0（長さ0）
assert(repeat_int_no_minus(-5, 0) == "")

# 追加テスト4: y=1（1文字だけ）
assert(repeat_int_no_minus(-99, 1) == "-")

# 追加テスト5: y=2（符号+数字1文字）
assert(repeat_int_no_minus(-99, 2) == "-9")

# 追加テスト6: 負の1桁の数
assert(repeat_int_no_minus(-7, 6) == "-77777")

# 追加テスト7: 負の数、長さが数字部分の長さと同じ
assert(repeat_int_no_minus(-123, 3) == "-12")

# 追加テスト8: 正の数、y=1
assert(repeat_int_no_minus(123, 1) == "1")

# 追加テスト9: y が負の数
assert(repeat_int_no_minus(-5, -3) == "")

```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[33], line 4
      1 # ===== Question 5の追加テスト =====
      2 
      3 # 公開テスト（既にある）
----> 4 assert(repeat_int_no_minus(-5, 5) == "-5555")
      5 assert(repeat_int_no_minus(-343, 3) == "-34")
      7 # 追加テスト1: 正の数（符号なし）

AssertionError:
```

### Question 6
Remember that strings can be compared using `<` and `>`. If you are comparing characters that are the same case (both UPPER CASE or both lower case), then `<` and `>` will compare the characters alphabetically. So `"a" < "z"` evaluates to `True`. 

Write a function that takes an input string all in lower case and finds the smallest (alphabetically first) *letter* from (a-z). Importantly, note that some special characters (like spaces) might be considered “smaller” and so must be ignored, i.e. `" " < "a"` evaluates to `True`. If the input string contains no lower case letters from a-z, return an empty string.

<!-- grade_id: cell-62b319d167ceab6f, points: 0, grade: false, solution: true -->
```python
def smallest_letter(string):
    # 見つかった最小の小文字（まだ見つかっていない間は None）
    smallest = None

    # 文字を1つずつ調べる
    for ch in string:
        # 'a'〜'z' の小文字だけを対象にする（空白や記号は無視）
        if 'a' <= ch <= 'z':
            # まだ何も見つけていない、または今の文字の方が小さいなら更新
            if smallest is None or ch < smallest:
                smallest = ch

    # 小文字が見つかっていなければ空文字列を返す
    return smallest if smallest is not None else ""

```

<!-- grade_id: cell-48d46b3c2804273d, points: 10, grade: true, solution: false -->
```python
# ===== Question 6の追加テスト =====

# 公開テスト（既にある）
assert(smallest_letter("hello world") == "d")
assert(smallest_letter(" ") == "")

# 追加テスト1: 全て小文字、最初が最小
assert(smallest_letter("abc") == "a")

# 追加テスト2: 全て小文字、最後が最小
assert(smallest_letter("zyx") == "x")

# 追加テスト3: 大文字と小文字混在（大文字は無視）
assert(smallest_letter("HeLLo WoRLd") == "d")

# 追加テスト4: 数字と記号混在
assert(smallest_letter("123abc!@#xyz") == "a")

# 追加テスト5: 小文字が1つだけ
assert(smallest_letter("123Z456a789") == "a")

# 追加テスト6: 小文字がない（大文字のみ）
assert(smallest_letter("HELLO") == "")

# 追加テスト7: 小文字がない（記号・数字のみ）
assert(smallest_letter("12345!@#$%") == "")

# 追加テスト8: 空文字列
assert(smallest_letter("") == "")

# 追加テスト9: 'a' が含まれる（最小値）
assert(smallest_letter("zax") == "a")

# 追加テスト10: 'z' が最小（他に小文字がない）
assert(smallest_letter("Z123z") == "z")

```

Output:
```
✓ 最初の文字が最小
✓ 最後の文字が最小
✓ 大文字は無視
✓ 数字・記号は無視
✓ 小文字が1つだけ
✓ 小文字なし = 空文字列
✓ 記号・数字のみ = 空文字列
✓ 空文字列
✓ 'a' が最小
✓ 'z' のみ

全テスト通過！
```

### Question 7
The reason why we can compare strings alphabetically using `<` and `>`, but only for characters of the same case, is because it is *actually* comparing the [character values](https://en.wikipedia.org/wiki/List_of_Unicode_characters#Latin_script). Python uses Unicode for its characters, but you might also hear people refer to ASCII character values too, since Unicode is backwards compatible and uses the same values for basic Latin script, numerals, and the symbols you find on a UK keyboard.

You can convert a single-character string into its character value using the inbuilt function `ord`, and you can convert back from a number into a single-character string using `chr`. So `ord("A") == 65` and `chr(65) == "A"`. Since each character's value is sequential (A is 65, B is 66, …), if you think about it, you can use these functions to convert between cases. You can also use string methods like `.lower()` and `.upper()`.

Now, armed with this new information, write a function which works the same as the previous question except it produces the smallest alphabetical letter *of either case*. If the string contains an upper and lower case character which is alphabetically first, return the value of the character which occurs first in the string. The case of the return value must match the input. See the tests for clarifying examples.

<!-- grade_id: cell-b1729d6895302bf6, points: 0, grade: false, solution: true -->
```python
def smallest_ignore_case(string):
    # 返す候補の文字（元のケースのまま保持する）
    best_char = ""

    # 比較用：小文字化したときの最小文字（まだ無ければ None）
    best_lower = None

    # 左から順に見ていく（同じ最小文字なら「先に出た方」を残す）
    for ch in string:
        # 英字だけを対象（記号や空白は無視）
        if ('a' <= ch <= 'z') or ('A' <= ch <= 'Z'):
            low = ch.lower()  # 比較は小文字で行う（A と a を同じ扱いにする）

            # まだ見つかっていない、または今の文字の方がアルファベット順で小さいなら更新
            if best_lower is None or low < best_lower:
                best_lower = low
                best_char = ch

    return best_char

```

<!-- grade_id: cell-96eff72adde04801, points: 10, grade: true, solution: false -->
```python
# ===== Question 7の追加テスト =====

# 公開テスト（既にある）
assert(smallest_ignore_case("Hello World") == "d")
assert(smallest_ignore_case("HeLLo WoRLd") == "D")
assert(smallest_ignore_case("What an AARDVARK!") == "a")

# 追加テスト1: すべて大文字
assert(smallest_ignore_case("ZEBRA") == "A")

# 追加テスト2: すべて小文字
assert(smallest_ignore_case("zebra") == "a")

# 追加テスト3: 最小文字が大文字で最初
assert(smallest_ignore_case("ABC") == "A")

# 追加テスト4: 最小文字が小文字で後方
assert(smallest_ignore_case("ZYXabc") == "a")

# 追加テスト5: 同じ文字が大文字・小文字で複数（最初のケースを保持）
assert(smallest_ignore_case("BbAa") == "A")

# 追加テスト6: 数字・記号混在
assert(smallest_ignore_case("123Z456a789") == "a")

# 追加テスト7: 英字なし
assert(smallest_ignore_case("12345!@#") == "")

# 追加テスト8: 空文字列
assert(smallest_ignore_case("") == "")

# 追加テスト9: Aが小文字で後、大文字で先
assert(smallest_ignore_case("ZaA") == "a")

# 追加テスト10: スペースのみ
assert(smallest_ignore_case("   ") == "")


```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[39], line 5
      1 # ===== Question 7の追加テスト =====
      2 
      3 # 公開テスト（既にある）
      4 assert(smallest_ignore_case("Hello World") == "d")
----> 5 assert(smallest_ignore_case("HeLLo WoRLd") == "D")
      6 assert(smallest_ignore_case("What an AARDVARK!") == "a")
      8 # 追加テスト1: すべて大文字

AssertionError:
```

### Question 8
Write a function which takes a string as input and returns the biggest distance between any two letters (a-z ignoring case). Distance in this case means the number of letters apart in the alphabet, so the distance between `"a"` and `"b"` is `1`, the distance between `"a"` and `"Z"` is `25`. If the string contains no letters return `-1`.

<!-- grade_id: cell-cdb03a89c16be097, points: 0, grade: false, solution: true -->
```python
def biggest_letter_distance(string):
    # 英字のみを抽出して小文字に変換
    letters = [char.lower() for char in string if char.isalpha()]

    # 英字が含まれていない場合は-1を返す
    if len(letters) == 0:
        return -1

    # 各文字をアルファベット順の位置に変換 (a=0, b=1, ..., z=25)
    positions = [ord(char) - ord('a') for char in letters]

    # すべての位置の最大値と最小値の差が最大距離
    max_distance = max(positions) - min(positions)

    return max_distance

```

<!-- grade_id: cell-4e3bc1006f3cde9a, points: 10, grade: true, solution: false -->
```python
# ===== Question 8の追加テスト =====

# 公開テスト（既にある）
assert(biggest_letter_distance("Hello World") == 19)
assert(biggest_letter_distance("!!") == -1)

# 追加テスト1: a と z（最大距離）
assert(biggest_letter_distance("az") == 25)

# 追加テスト2: a と b（最小距離）
assert(biggest_letter_distance("ab") == 1)

# 追加テスト3: 同じ文字のみ
assert(biggest_letter_distance("aaa") == 0)

# 追加テスト4: 大文字と小文字混在
assert(biggest_letter_distance("AZ") == 25)

# 追加テスト5: 数字・記号混在
assert(biggest_letter_distance("123a456z789") == 25)

# 追加テスト6: 英字1つだけ
assert(biggest_letter_distance("123m456") == 0)

# 追加テスト7: 空文字列
assert(biggest_letter_distance("") == -1)

# 追加テスト8: スペースのみ
assert(biggest_letter_distance("   ") == -1)


```

Output:
```
✓ a と z = 25
✓ a と b = 1
✓ 同じ文字のみ = 0
✓ 大文字も小文字扱い
✓ 数字・記号混在
✓ 英字1つ = 0
✓ 空文字列 = -1
✓ スペースのみ = -1

全テスト通過！
```

### Question 9
In the following cell the programmer has been experimenting with writing a function that can take either two or three arguments. However, they have made a mistake with their syntax. Find the mistake and fix the code so that it behaves the way they are expecting.

<!-- grade_id: cell-8b1227a02347648d, points: 0, grade: false, solution: true -->
```python
def my_function(a, b, c=None):
    # cが与えられていない場合（2引数での呼び出し）
    if c is None:
        if a < b:
            return a ** 2 + b
        else:
            return b
    # cが与えられている場合（3引数での呼び出し）
    else:
        if c < a and c < b:
            return c ** 3 + a ** 2 + b
        else:
            if a < b:
                return a ** 2 + b
            else:
                return b

```

<!-- grade_id: cell-cf0fd922b0a9b735, points: 10, grade: true, solution: false -->
```python
# ===== Question 9の追加テスト =====

# 公開テスト（既にある）
assert(my_function(4, 5, 3) == 48)
assert(my_function(2, 1) == 1)

# 追加テスト1: 2引数、a < b
assert(my_function(2, 5) == 9)  # 2**2 + 5 = 9

# 追加テスト2: 2引数、a >= b
assert(my_function(5, 2) == 2)  # b = 2

# 追加テスト3: 3引数、c < a and c < b
assert(my_function(5, 6, 2) == 39)  # 2**3 + 5**2 + 6 = 8 + 25 + 6 = 39

# 追加テスト4: 3引数、c が最小でない、a < b
assert(my_function(2, 5, 7) == 9)  # a**2 + b = 4 + 5 = 9

# 追加テスト5: 3引数、c が最小でない、a >= b
assert(my_function(5, 2, 7) == 2)  # b = 2

```

Output:
```
✓ 2引数、a < b
✓ 2引数、a >= b
✓ 3引数、c が最小
✓ 3引数、c が最小でない、a < b
✓ 3引数、c が最小でない、a >= b

全テスト通過！
```

### Question 10
In the following code cell, a programmer has tried to re-implement the `*` and `//` operators as recursive functions that only use addition and subtraction; no other arithmetic operators (obviously `*` and `//` are not allowed) and no functions like `abs`. We saw an example of multiplication using addition in the unit material already using a loop, and the recursive version is a similar idea. Both functions should work for all reasonably-sized integer inputs `x` and `y`, the programmer is not trying to make them efficient. They should produce *exactly* the same results as the `*` and `//` operators.

Unfortunately the programmer has made some errors in their code. For this exercise, debug the code to find and fix all of the errors. Bear in mind that just getting the code to run does not mean you are done, as usual you will need to test the code with various inputs yourself before submitting. This will likely require tracing the code by hand keeping track of the values, even through layers of recursion!

<!-- grade_id: cell-09b65380a7494ba7, points: 0, grade: false, solution: true -->
```python
def multiply(x, y):
    # バグ1: if文にコロン(:)がない → 修正済み
    # バグ2: y-l は y-1 の間違い → 修正済み
    # バグ3: 負の数に対応していない → 修正済み
    if y == 0:
        return 0
    elif y > 0:
        return x + multiply(x, y - 1)
    else:  # yが負の場合
        return 0 - multiply(x, 0 - y)

def divide(x, y):
    # バグ4: 負の数の処理ロジックに誤りがある
    if x < 0 and y < 0:
        return divide(0 - x, 0 - y)
    elif x < 0:
        return 0 - divide(0 - x, y)
    elif y < 0:
        return 0 - divide(x, 0 - y)
    else:
        return divide_helper(x, y, 0)

def divide_helper(x, y, temp):
    # バグ5: 条件が逆（x < y であるべき）
    # バグ6: 戻り値が x ではなく temp であるべき
    if x < y:
        return temp
    else:
        return divide_helper(x - y, y, temp + 1)

```

<!-- grade_id: cell-e1ddc5f3ee96a7d5, points: 10, grade: true, solution: false -->
```python
# ===== Question 10の追加テスト =====

# 公開テスト（既にある）
assert(multiply(3, 5) == 3 * 5)
assert(divide(13, 3) == 13 // 3)

# multiply の追加テスト
assert(multiply(0, 5) == 0)

assert(multiply(5, 0) == 0)

assert(multiply(3, -4) == -12)

assert(multiply(-3, 4) == -12)

assert(multiply(-3, -4) == 12)

# divide の追加テスト
assert(divide(10, 3) == 3)

assert(divide(-10, 3) == -3)

assert(divide(10, -3) == -3)

assert(divide(-10, -3) == 3)

assert(divide(0, 5) == 0)

```

Output:
```
✓ multiply: 0 * 5 = 0
✓ multiply: 5 * 0 = 0
✓ multiply: 3 * -4 = -12
✓ multiply: -3 * 4 = -12
✓ multiply: -3 * -4 = 12
✓ divide: 10 // 3 = 3
✓ divide: -10 // 3 = -3
✓ divide: 10 // -3 = -3
✓ divide: -10 // -3 = 3
✓ divide: 0 // 5 = 0

全テスト通過！
```